```python
NAME = "Jefferson"
COLLABORATORS = "None"
```

---

# CSE 20 Testout Sample

There are 6 questions, each worth 10 points. For each question, there is both a description of the questions, and some tests that help you check that your solution works.  Some questions have two sets of tests, easier ones, and full tests, so that you can get some points even if you do not solve the full problem. 

*   Try out your own test cases. Even if your code work on the given test cases, that does not mean you will get full credit. We have hidden test cases that are not shown here. So please test your code.
*   When inputs are lists: common errors happen when either the input is empty or has one element. Make sure you test such cases.
*   You can write out print statements in your code, to help you test/debug it.

Once you are done, if this were a real exam, you would submit your work as follows:

* Download the notebook via File > Save as .ipynb
* Upload it to a specified google form.


## Preliminaries

First, let us check that you are working with Python 3, and that you have the package required for the tests.

**Nose.** The line below works on colab.research.google.com.  If it does not work on your laptop, do:

    !pip install nose
    
if you use pip to install packages, and 

    !conda install nose
    
if you use conda

```python
!pip install nose
from nose.tools import assert_equal, assert_true, assert_false
```

Output:
```
Looking in indexes: https://pypi.org/simple, https://us-python.pkg.dev/colab-wheels/public/simple/
Collecting nose
  Downloading nose-1.3.7-py3-none-any.whl (154 kB)
[K     |████████████████████████████████| 154 kB 14.5 MB/s 
[?25hInstalling collected packages: nose
Successfully installed nose-1.3.7
```

```python
# We check that you are running Python 3.  You need to use Python 3.
import sys
assert_equal(sys.version_info.major, 3)
```

## Question 1

Write a function ziplist that, given two lists, produces a list of pairs, with the first element of the pair coming from the first list, and the second from the second list.  The output list of pairs must be as long as the longest of the input lists, and the missing elements from the shorter input list give rise to None in the output list. 


<!-- grade_id: cell-9d983671ab005c31, points: 0, grade: false, solution: true -->
```python
def ziplist(a,b):
  # YOUR CODE HERE
    if len(a)>len(b):
      x = len(a)-len(b)
      for t in range(x):
        b.append(None)
    elif len(b)>len(a):
      e=len(b)-len(a)
      for s in range(e):
        a.append(None)
    p = []
    h =0
    for k in a:
      y = (k,b[h])
      p.append(y)
      h+=1
    return p

```

explanation: My main thought was just to make a for loop that appends a value from list a, then list b and so on. Before that I needed the lists to be equal length, and as the questions references, if one list is longer than the other, then the unpaired values should be paired with a None value. So the first if/elif portion of my code if figuring out which list is longer and adding enough None values to the shorter list to make them equal length. Then finally, I just made a for-loop that added on the pairs.

```python
ziplist(['a','b','c'],[1,2,3])
```

Output:
```
[('a', 1), ('b', 2), ('c', 3)]
```

<!-- grade_id: cell-2e2fe7596f2b2774, points: 4, grade: true, solution: false -->
```python
# 4 Points. First, let's check that your code works on lists of equal length.
assert_equal(ziplist(['a', 'b', 'c'], [1, 2, 3]), [('a', 1), ('b', 2), ('c', 3)])
assert_equal(ziplist([], []), [])
```

<!-- grade_id: cell-34ebe737960135e8, points: 6, grade: true, solution: false -->
```python
# 6 Points. The function also needs to work on lists of different lengths.
assert_equal(ziplist([1, 2], []), [(1, None), (2, None)])
assert_equal(ziplist([1], ['a', 'b']), [(1, 'a'), (None, 'b')])
```

## Question 2

Write a function threeormore that, given an input list, outputs the set of list elements that occur 3 or more times.



<!-- grade_id: cell-60ac8898a224d9c8, points: 0, grade: false, solution: true -->
```python
def threeormore(li):
  # YOUR CODE HERE
    end = set([])
    cuntr=0
    for key in li:
      for inkey in li:
          if key==inkey:
            cuntr+=1
      if cuntr>=3:
        end.add(key)
      cuntr=0
    return end
```

explanation: Since you're checking if something the lists occurs in the rest of the list 3+ times, then my thought was to make a nested for-loop. The outer for-loop starts at index 0, and the inner for-loop checks the entire list if the outer key == the inkey, and if so then the cuntr goes up 1. If it occurs 3+ times then the key gets added to the set. I used a set because if the key is already listed in the set then it simply doesn't get added again, stopping repeating numbers.

```python
threeormore([1,2,3,2,2])
```

Output:
```
{2}
```

<!-- grade_id: cell-57eca55510d54862, points: 10, grade: true, solution: false -->
```python
assert_equal(threeormore([1, 3, 5, 4, 3, 7, 6, 3, 7, 7, 7, 1, 2, 0, 3, 1]), {1, 3, 7})
assert_equal(threeormore([]), set())
```

## Question 3

Implement a queue class: a class MyQueue with two operations, enqueue and dequeue. A queue takes in items in first-in-first-out order. The enqueue function adds items to a queue. The dequeue function returns the earliest item enqueued (and deletes this from the queue). If the queue is empty, then dequeue should return None.

You can look at a specific example below.

<!-- grade_id: cell-36c0674aaa538522, points: 0, grade: false, solution: true -->
```python
class MyQueue(): 
    # YOUR CODE HERE
    def __init__(self,q=[]):
        self.q=q     
    def enqueue(self, x):
        self.q.append(x)
    def dequeue(self):
        if len(self.q)>0:
          return self.q.pop(0) 
        else:
          return None
```

explanation: enqueue is just adding onto the list. for the dequeue function, I used the pop method as you can return the earliest value while also removing it.

```python
b=MyQueue()
b.enqueue('banana')
b.enqueue('apple')
b.dequeue()
b.enqueue('cherry')
b.dequeue()
b.dequeue()
b.dequeue()
```

<!-- grade_id: cell-2e84ce35e27cbe8b, points: 10, grade: true, solution: false -->
```python
# Example of behavior.
q = MyQueue()
q.enqueue('banana')
q.enqueue('apple')
assert_equal(q.dequeue(),'banana')
q.enqueue('cherry')
assert_equal(q.dequeue(), 'apple')
assert_equal(q.dequeue(), 'cherry')
assert_equal(q.dequeue(), None)
```

## Question 4

Implement a function closest_neighbor that, given a list of points in the (x, y) format, returns the pair of closest points (according to the usual notion of distance of points on the plane), listing the pair in the same relative order as in the list. If the list has fewer than two elements, return None.

<!-- grade_id: cell-ca3c2aecc698ef1b, points: 0, grade: false, solution: true -->
```python
import math
def closest_neighbor(point_list):
    # YOUR CODE HERE
    leng = len(point_list)
    if len(point_list)<=2:
      return None
    else:
      k={}
      for key in point_list[0:leng-1]:
        x=0
        for inkey in point_list[x+1:]:
          dist = float(math.sqrt((inkey[0]-key[0])**2+(inkey[1]-key[1])**2))
          if(key!=inkey):
            k[key,inkey]=float(dist)
          else:
            pass
        x+=1
    val=2**31-1
    fay=()
    for fin in k:
      if k[fin]<val:
        fay=fin
        val=k[fin]
    return fay
```

notes: 
Someone in the ucsc 2026 discord said it's discouraged to import anything, but I didn't want to go through the trouble of making a sqrt method
Also this question was the one that gave me the most trouble and is most likely not 100% correct

explanation: As the question states if the list has less than 2 elements then return None. Like question 2, I made a nested for-loop. Then for every key(pair), the inner for-loop would go through every pair after the original key and assign it a distance. This would then be placed into the dictionary (k) I made. This would in turn list every possible combination of coordinate pairs along with it's value(distance). Then finally, I just created a for loop to go through the values to find the smallest dist. and I returned the key for that dist.


```python
closest_neighbor([(1,2),(5,6),(3,9)])
```

Output:
```
((5, 6), (3, 9))
```

<!-- grade_id: cell-dc8b60881d9ca1c5, points: 10, grade: true, solution: false -->
```python


assert_equal(closest_neighbor([(1, 2), (4, 5), (3,1), (5,5)]), ((4, 5), (5, 5)))

assert_equal(closest_neighbor([(1, 2)]), None)
```

## Question 5

Implement a function olympic_average that, given a list of numbers, if the list is of length at least 3, excludes the minimum and maximum elements and returns the average of the other elements. If the list is shorter than 3, it just returns the average of the elements. If the list is empty, it returns None.

<!-- grade_id: cell-f3731568c71eecc8, points: 0, grade: false, solution: true -->
```python
def olympic_average(els):
    # YOUR CODE HERE
    els.sort()
    chefc=0
    denom=0
    if len(els)>=3:
      for num in els[1:len(els)-1]:
        chefc+=num
        denom+=1
      return float(chefc/denom)
    elif len(els)!=0:
      for num in els:
        chefc+=num
        denom+=1
      return float(chefc/denom)
    return None
```

explanation: I first started for cases in lists that had 3+ elements, so I made a for-loop that ran through index's 1 through len(list)-1. And then it's simply finding an avg. My elif were for lists less than 3 elements and is also just finding an avg. Finally return None if there's no elements in given list.

<!-- grade_id: cell-8145a59482427824, points: 5, grade: true, solution: false -->
```python
# 5 Points: Let us assume that the list has at least 3 elements.
assert_equal(olympic_average([5, 3, 4, 5]), 4.5)
assert_equal(olympic_average([5, 6, 3, 4, 10]), 5)
```

<!-- grade_id: cell-53ffd7d4c2662095, points: 5, grade: true, solution: false -->
```python
# 5 Points: General case.
assert_equal(olympic_average([2, 3]), 2.5)
assert_equal(olympic_average([4]), 4)
assert_equal(olympic_average([]), None)
assert_equal(olympic_average([5, 3, 11, 5]), 5)
```

## Question 6

Implement a function is_anagram that takes as input two strings, and returns True if and only if the strings are anagrams of each other. This means that the both strings are permutations of each other. You can assume that string will only have lowercase letters.

<!-- grade_id: cell-166fd5fa32f2c27f, points: 0, grade: false, solution: true -->
```python
def is_anagram(w1, w2):
    """Returns True iff word w1 is an anagram of word w2."""
    # YOUR CODE HERE
    word1 = list(w1)
    word2= list(w2)
    word1.sort()
    word2.sort()
    if word1==word2:
      return True
    else:
      return False
```

explanation: If 2 words are anagrams of each other, then them being sorted in a list should equal to each other. so I turned the strings into lists, then sorted them, then checked if they were equivalent.

<!-- grade_id: cell-19398d91ea6ebef7, points: 5, grade: true, solution: false -->
```python
# 5 points. First, let's assume that the word does not contain repeated letters. 
assert_true(is_anagram("home", "hemo"))
assert_true(is_anagram("", ""))
assert_true(is_anagram("figure", "refugi"))
assert_false(is_anagram("seal", "coal"))
```

<!-- grade_id: cell-59f0616d30eafe4a, points: 5, grade: true, solution: false -->
```python
# 5 points.  General case, where there can be repeated characters. 
assert_true(is_anagram("banana", "ananab"))
assert_false(is_anagram("hello", "helo"))
assert_false(is_anagram("banana", "bnana"))
```